<subject>
Assignment How-To
</subject>

<details>
    
**Do Not Talk During Quizzes or Exams**

Do not talk at all once an exam has begun. You may talk again once you leave the room after the quiz or exam. The one exception to this rule is if you need to ask the exam proctor a question. Any talking during a quiz or exam qualifies the student for disciplinary action.

**Naming Conventions**

When naming your files for upload, you must follow the format below:

    <uni>_<assignment>_<details [optional]>.<extension>

For example, if I needed to hand in HW 0, any of the below formats would be sufficient for a file name:

- pl2648_hw0.ipynb
- pl2648_hw0.txt
- pl2648_hw0.sh
- pl2648_hw0_all_in_one.txt
- Pl2648_hw0_bash_program.sh
- Pl2648_quiz1.ipynb

This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

Courseworks may rename your file to something like `pl2648_hw0-1.ipynb` if you resubmit your assignment. This is perfectly fine.

**What Format To Submit In**

Most homework and quizzes are in Jupyter notebooks. Unless specified otherwise, please download your work as an `.ipynb` file from your local machine and upload it to courseworks.


**Grading**

Possible points on late assignments are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors. 

Because grading is automated, please delete (or comment out) the `raise NotImplmeneted` code before attempting a problem. 

Empty un-editable cells in an assignment are there for a purpose. They will be filled with tests by the automatic grader. Please do not attempt to remove them.

**Getting Help**

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. You can not post screenshots of your current work to Piazza or other tools used for getting help.

If you need to reach out to a CA for help, please do so via Piazza and not via email. Answers given via Piazza will help you as well as other students. Thus, emails will always have a lower priority for response than Piazza questions. If you do email the CA, please make a note of what section you are in. This helps us identify you in courseworks faster. 

Finally, if you do not get a repsonse from a CA within 48 hours, you may email the professor.

**Multiple Choice**

If the question is multiple choice, you will be given several options to choose from and your function will need to return **one** of those options **verbatim** as a string.

For example: 

Which of the following animals bark?

- dogs
- cats
- fish
- trees

A correct answer would be structured in the following way:

```python
def question_animals():
    return 'dogs'
```

You answer will be stripped of left and right white space and lowercased before comparison to the correct answer during grading.
</details>

<!-- grade_id: cell-e5713946811b5d52, points: 0, grade: true, solution: false -->
```python
### BEGIN TESTS
import subprocess
import tempfile
import os


def clean_bash_function(text):
    if 'rm' in text:
        raise ValueError('rm found in text of script. No points.')
    
    text = text.strip()
    # Get first \n
    lines = text.split('\n')

    return '\n'.join(line.lstrip() for line in lines)


def run_script(text, args=None, timeout=None):
    args = args or []
    text = clean_bash_function(text)
    with tempfile.NamedTemporaryFile() as file:
        file.write(text.encode())
        file.flush()
        os.chmod(file.name, 0o777)    
        response = subprocess.run(
            ['/bin/bash', file.name] + args, 
            timeout=timeout, 
            stdout=subprocess.PIPE, 
            stderr=subprocess.PIPE,
        )
        return response
### END TESTS
```

**DO NOT INCLUDE THE STRING  `rm`  ANYWHERE IN YOUR SCRIPTS OR YOU COULD GET ZERO POINTS FOR THE QUESTION.**

**Q**: Write a function `get_bash_function` that returns the text of a bash script. For example:

```python
def get_bash_function():
    return """
    #!/bin/bash
    echo Hello!
    """
```

The bash script should not return "hello" but should count down from 5 to 1 (inclusive), printing each number as it goes and should print "GO!" after printing 1. All values should be printed to their own line. All values should be printed to standard out. The script should pause for one second in between printing each line. A `for` loop should be used in the count down. If you script were run at the command line, output would be this:

```bash
$ /bin/bash your-script.sh
5
4
3
2
1
GO!
```

[2 points]

<!-- grade_id: cell-40a07e3e81c26639, points: 0, grade: false, solution: true -->
```python
def get_bash_function():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ae95381e40692307, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
text = get_bash_function()
response = run_script(text, timeout=10)
assert response.stdout.decode().strip().split() == ['5', '4', '3', '2', '1', 'GO!']
### END TESTS
```

<!-- grade_id: cell-03858a9510d75940, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-c26d6af528775061, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-de51c6a3de2b41e7, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-396d56170dc6809e, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-fb1f7cddbdd9b665, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-e562b568261ce8ae, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-85db33b0a319d4c7, points: 0.25, grade: true, solution: false -->
```python

```

**Q**: Write a function `get_bash_function` that returns the text of a bash script. For example:

```python
def get_bash_function():
    return """
    #!/bin/bash
    echo Hello!
    """
```

The bash script should not return "hello" but should take one argument and use it within a greeting. Please use the greeting seen in the example. All output should be printed to standard out. If you script were run at the command line, output would be this:

```bash
$ /bin/bash your-script.sh Paul
Well, hello Paul! Nice to meet you today :)
```

This is a simple problem. It's designed to ensure you know how to write a bash script that takes command line arguments.

[1 point]

<!-- grade_id: cell-0ae8eeb9e364516c, points: 0, grade: false, solution: true -->
```python
def get_bash_function():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2af32a0dc3f329d1, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-45b88ca419f900fc, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-8651e4fc2599fb6b, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-8703911f31126e50, points: 0.25, grade: true, solution: false -->
```python

```

**Q**: Write a function `get_bash_function` that returns the text of a bash script. For example:

```python
def get_bash_function():
    return """
    #!/bin/bash
    echo Hello!
    """
```

The bash script should not return "hello" but should take one argument and find that argument's factorial. All output should be printed to standard out. If you script were run at the command line, output would be this:

```bash
$ /bin/bash your-script.sh 5
120
```
[2 points]

<!-- grade_id: cell-fd417f3e5c08b2a1, points: 0, grade: false, solution: true -->
```python
def get_bash_function():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-3b22872326e5c128, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
text = get_bash_function()

response = run_script(text, args=['5'], timeout=2)

assert response.stdout.decode().strip() == '120'
### END TESTS
```

<!-- grade_id: cell-cd0657b68cceb4f2, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-3c0ceafbf4651462, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-bd91b931ff931cca, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
text = get_bash_function()

response = run_script(text, args=['0'], timeout=2)

assert response.stdout.decode().strip() == '1'
### END TESTS
```

<!-- grade_id: cell-7528b05a98e71db9, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-e14c5cb2635de833, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-dc839d76dd4f9ee5, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-bde2de44e3a104b1, points: 0.25, grade: true, solution: false -->
```python

```